# Module 1

#### In this assignment, you will work with ufo sightings data.
- The data includes various data points about individual ufo sightings
- Data File(s): ufo-sightings.csv

<!-- grade_id: init, points: 0, grade: true, solution: false -->
```python
###########################################################
### EXECUTE THIS CELL BEFORE YOU TO TEST YOUR SOLUTIONS ###
###########################################################

import imp, os, sys
sol = imp.load_compiled("solutions", "./solutions.py")
sol.get_solutions("ufo-sightings.csv")
from nose.tools import assert_equal
```

<!-- grade_id: q1, points: 0, grade: false, solution: true -->
```python
'''
1. Import the csv module. Load and read the UFO sightings data set, from the ufo-sightings.csv file, 
into a DictReader inside a with statement.  Assume the data file is in the same directory as the code. 

Print the field names of the data set. Iterate over the reader to put the data into a list name "ufosightings".

'''

import csv
filepath = "ufo-sightings.csv"
ufosightings = [] 

# your code here
with open(filepath, 'r') as csvfile:
    reader = csv.DictReader(csvfile)
    n = 0
    for row in reader:
        ufosightings.append(row)
        n +=1


```

```python
print(n)
ufosightings[0]
```

Output:
```
80332
OrderedDict([('datetime', '10/10/1949 20:30'),
             ('city', 'san marcos'),
             ('state', 'tx'),
             ('country', 'us'),
             ('shape', 'cylinder'),
             ('duration (seconds)', '2700'),
             ('duration (hours/min)', '45 minutes'),
             ('comments',
              'This event took place in early fall around 1949-50. It occurred after a Boy Scout meeting in the Baptist Church. The Baptist Church sit'),
             ('date posted', '4/27/2004'),
             ('latitude', '29.8830556'),
             ('longitude ', '-97.9411111')])
```

<!-- grade_id: q1_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(ufosightings, sol.ufosightings)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q2, points: 0, grade: false, solution: true -->
```python
'''
2. How many sightings were there in total?  Put the count in "ufosightings_count" and print the result.
'''
# your code here
ufosightings_count = 0
for i in ufosightings:
    ufosightings_count+=1

print(ufosightings_count)
```

Output:
```
80332
```

<!-- grade_id: q2_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(ufosightings_count, sol.ufosightings_count)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q3, points: 0, grade: false, solution: true -->
```python
'''
3. How many sightings were there in the US?  Put the US sightings in "sightings_us" and print.

Hint: Check for ufo sightings where the country is 'us'.

'''

# your code here
sightings_us = 0
sightings_us = [i for i in ufosightings if i["country"] == "us"]
sightings_us
```

Output:
```
[OrderedDict([('datetime', '10/10/1949 20:30'),
              ('city', 'san marcos'),
              ('state', 'tx'),
              ('country', 'us'),
              ('shape', 'cylinder'),
              ('duration (seconds)', '2700'),
              ('duration (hours/min)', '45 minutes'),
              ('comments',
               'This event took place in early fall around 1949-50. It occurred after a Boy Scout meeting in the Baptist Church. The Baptist Church sit'),
              ('date posted', '4/27/2004'),
              ('latitude', '29.8830556'),
              ('longitude ', '-97.9411111')]),
 OrderedDict([('datetime', '10/10/1956 21:00'),
              ('city', 'edna'),
              ('state', 'tx'),
              ('country', 'us'),
              ('shape', 'circle'),
              ('duration (seconds)', '20'),
              ('duration (hours/min)', '1/2 hour'),
              ('comments',
               'My older brother and twin sister were leaving the only Edna theater at about 9 P
```

```python
# base iteration counter
n = 0
for i in sightings_us:
    n += 1
n
```

Output:
```
65114
```

```python
len(sightings_us) # len() of Dict
```

Output:
```
65114
```

```python
# count with lambda 
n = lambda x: sum(1 for i in sightings_us)
n(sightings_us)
```

Output:
```
65114
```

<!-- grade_id: q3_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(sightings_us,  sol.sightings_us)
print("Success!")
```

Output:
```
Success!
```

```python
sightings_us[5]
```

Output:
```
OrderedDict([('datetime', '10/10/1966 20:00'),
             ('city', 'pell city'),
             ('state', 'al'),
             ('country', 'us'),
             ('shape', 'disk'),
             ('duration (seconds)', '180'),
             ('duration (hours/min)', '3  minutes'),
             ('comments',
              'Strobe Lighted disk shape object observed close&#44 at low speeds&#44 and low altitude in Oct 1966 in Pell City Alabama'),
             ('date posted', '3/19/2009'),
             ('latitude', '33.5861111'),
             ('longitude ', '-86.2861111')])
```

<!-- grade_id: q4, points: 0, grade: false, solution: true -->
```python
'''
4. Let's find the "fireball" sighting(s) that lasted more than ten seconds in US. 
Print the the datetime and state of each.  Put the data in "fball" and print the result.

Note: Consider only the US sightings stored in "sightings_us".

- Cast the duration in seconds to a float (decimal). 
- Check if duration is greater than 10. 
- Check if the shape is "fireball".

'''

#First, define a Python function that checks if a given duration (seconds) is "valid"
def is_valid_duration(duration_as_string):
# your code here
    try: duration_as_float = float(duration_as_string)
    except ValueError:
        return False
    else: return duration_as_float
    
fball = [ i for i in sightings_us if (i["shape"] == "fireball" and is_valid_duration(i["duration (seconds)"]) > 10) ]
fball
```

Output:
```
[OrderedDict([('datetime', '10/10/1968 19:00'),
              ('city', 'brevard'),
              ('state', 'nc'),
              ('country', 'us'),
              ('shape', 'fireball'),
              ('duration (seconds)', '180'),
              ('duration (hours/min)', '3 minutes'),
              ('comments',
               'silent red /orange mass of energy floated by three of us in western North Carolina in the 60s'),
              ('date posted', '6/12/2008'),
              ('latitude', '35.2333333'),
              ('longitude ', '-82.7344444')]),
 OrderedDict([('datetime', '10/10/1998 03:30'),
              ('city', 'phoenix (north part)'),
              ('state', 'az'),
              ('country', 'us'),
              ('shape', 'fireball'),
              ('duration (seconds)', '15'),
              ('duration (hours/min)', '10 to 15 sec'),
              ('comments',
               'bright green moving north to north east. it was bright green and kinda clear. quite fast too.'),
        
```

```python
len(fball)
```

Output:
```
3730
```

<!-- grade_id: q4_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(fball, sol.fball)
print("Success!")
```

Output:
```
Success!
```

```python
# list_years = [ row["Year"] for row in album_list if is_year_int(row["Year"]) ]
# list_years
# min(list_years)
# albums_sorted = sorted(albums, key = lambda x: is_valid_duration(x["Artist"]) )
```

<!-- grade_id: q5, points: 0, grade: false, solution: true -->
```python
'''
5. Sort the above list by duration. What was the datetime and duration of the longest sighting?  
Put the sorted list in "fballsorted" and print the result.

- Cast the duration in seconds to a float (decimal). 
- Sort in reverse order.

'''

# your code here
fballsorted = sorted(fball, key = lambda x: is_valid_duration(x["duration (seconds)"]), reverse=True )

# print(fballsorted)
fballsorted
```

Output:
```
[OrderedDict([('datetime', '8/3/2008 21:00'),
              ('city', 'virginia beach'),
              ('state', 'va'),
              ('country', 'us'),
              ('shape', 'fireball'),
              ('duration (seconds)', '10526400'),
              ('duration (hours/min)', '4 months'),
              ('comments',
               'this object was very high up and emmited no sound.&#33'),
              ('date posted', '8/12/2008'),
              ('latitude', '36.8527778'),
              ('longitude ', '-75.9783333')]),
 OrderedDict([('datetime', '12/9/1965 24:00'),
              ('city', 'newtown square'),
              ('state', 'pa'),
              ('country', 'us'),
              ('shape', 'fireball'),
              ('duration (seconds)', '6120000'),
              ('duration (hours/min)', '1700hrs'),
              ('comments',
               'DEC 1965  NEWTOWN SQUARE PENN   FLAMING BURNING OBJECT   MOVING ACROSS THE SKY VERY LOW AND SLOW'),
              ('date posted', '7/16/2006')
```

<!-- grade_id: q5_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(fballsorted, sol.fballsorted)
print("Success!")
```

Output:
```
Success!
```

```python
fball_max_duration = max(fball, key = lambda x: is_valid_duration(x["duration (seconds)"]) )
state = fball_max_duration['state']
state
```

Output:
```
'va'
```

<!-- grade_id: q6, points: 0, grade: false, solution: true -->
```python
'''
6. What state had the longest lasting "fireball"?   Put the state in "state" and print the result.

- Check if the shape is "fireball".
- Cast the duration in seconds to a float (decimal).
- Get the record with the largest (max) duration in seconds.
- Get the state for the record.

'''

# your code here
state = max(fball, key = lambda x: is_valid_duration(x["duration (seconds)"]) )['state']
print(state)
```

Output:
```
va
```

<!-- grade_id: q6_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(state, sol.state)
print("Success!")
```

Output:
```
Success!
```

```python

ufosightings_wo_0_duration = [ i for i in ufosightings if is_valid_duration(i["duration (seconds)"]) != 0  ]
min_duration = float(min(ufosightings_wo_0_duration, key = lambda x: is_valid_duration(x["duration (seconds)"]) )["duration (seconds)"])
# min_duration = float(min_duration)

```

<!-- grade_id: q7, points: 0, grade: false, solution: true -->
```python
'''
7. Let's assume that any sighting (of any shape) of 0 seconds is insignificant. 
Write code to filter out these extraneous records and get the shortest sighting overall now.  
Put the minimum duration in "min_duration" and print the result.  
Use ufosightings
Note: Consider all sightings stored in "ufosightings".

'''

# your code here
ufosightings_wo_0_duration = [ i for i in ufosightings if is_valid_duration(i["duration (seconds)"]) != 0  ]
min_duration = float(min(ufosightings_wo_0_duration, key = lambda x: is_valid_duration(x["duration (seconds)"]) )["duration (seconds)"])


```

<!-- grade_id: q7_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(min_duration, sol.min_duration)
print("Success!")
```

Output:
```
Success!
```

```python
sightings_shapes = [ i["shape"] for i in ufosightings ]
len(sightings_shapes)
```

Output:
```
80332
```

<!-- grade_id: q8, points: 0, grade: false, solution: true -->
```python
'''
8. What are the top 3 shapes sighted, and how many sightings were there for each? 

Note: Consider all sightings stored in "ufosightings".

- Create a new list "sightings_shapes" containing values from the "shape" column in ufosightings.  
- Create a new dictionary "count" with values of that column as keys and the counts as values.
- Get a list of the dictionary keys and values using the items() method.  This will return a list of key:value pairs.
Sort the list of key:value pairs in reverse order, from greatest (most sightings) to least.

Get the top 3 and store in "top3shapes".  Print the result.

'''

#Create a new list containing values from the "shape" column in ufosightings.
# your code here
sightings_shapes = [ i["shape"] for i in ufosightings ]



#Create a new dictionary with values of that column as keys and the counts as values.
# your code here
count = {}
for i in set(sightings_shapes):
    n = 0
    for j in sightings_shapes:
        if i == j:
            n += 1     
    count.update({i:n})
    

#Get a list of the dictionary keys and values (use the items() method) and sort them in reverse order, from greatest (most sightings) to least.
#Get and print the top 3.
# your code here

top3shapes = sorted(count.items(), key = lambda x: x[1] , reverse=True)[:3]
print(top3shapes)
```

Output:
```
[('light', 16565), ('triangle', 7865), ('circle', 7608)]
```

<!-- grade_id: q8_test_a, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(sightings_shapes, sol.sightings_shapes)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q8_test_b, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(count, sol.count)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q8_test_c, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(top3shapes, sol.top3shapes)
print("Success!")
```

Output:
```
Success!
```

```python

```